# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
import sklearn
print('sklearn: %s' % sklearn.__version__)
iris = datasets.load_iris()
#raise NotImplementedError()
```

Output:
```
sklearn: 0.21.2
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf", gamma="auto")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
#raise NotImplementedError()
X = iris.data
y = iris.target

classifier = SVC(kernel="rbf", gamma="auto")
classifier.fit(X,y)
predictions = classifier.predict(X)
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```

```python
print(X[:5])
```

Output:
```
[[5.1 3.5 1.4 0.2]
 [4.9 3.  1.4 0.2]
 [4.7 3.2 1.3 0.2]
 [4.6 3.1 1.5 0.2]
 [5.  3.6 1.4 0.2]]
```

```python
len(y)
```

Output:
```
150
```

```python
len(X)
```

Output:
```
150
```

```python
list(X[:5])
```

Output:
```
[array([5.1, 3.5, 1.4, 0.2]),
 array([4.9, 3. , 1.4, 0.2]),
 array([4.7, 3.2, 1.3, 0.2]),
 array([4.6, 3.1, 1.5, 0.2]),
 array([5. , 3.6, 1.4, 0.2])]
```

```python
list(y[:5])
```

Output:
```
[0, 0, 0, 0, 0]
```

```python
list(iris)
```

Output:
```
['data', 'target', 'target_names', 'DESCR', 'feature_names', 'filename']
```

```python
list(iris.target_names)
```

Output:
```
['setosa', 'versicolor', 'virginica']
```

```python
list(iris.feature_names)
```

Output:
```
['sepal length (cm)',
 'sepal width (cm)',
 'petal length (cm)',
 'petal width (cm)']
```

```python
print(iris.filename)
```

Output:
```
/Users/rdoraiswamy/anaconda3/lib/python3.7/site-packages/sklearn/datasets/data/iris.csv
```

```python

```

```python

```